---
title: 'Descriptive Statistics'
description: 'This is a template chapter.'
---

## An exercise title written in sentence case

```yaml
type: NormalExercise
key: f7737c1239
lang: r
xp: 100
skills: 1
```

This is not really a Datacamp exercise. We will run a short experiment and analyze the sample we collect.

`@instructions`
- [Participate](http://www.onlinestrooptest.com/stroop_effect_test.php) in an online Stroop test.
- Open the [Google sheet](https://docs.google.com/spreadsheets/d/1IsyrFoKTCBpAOPnkdrwmHZnse6psvvZ1HCXyIaV50Jo/edit?usp=sharing) we will use to track you results.
- Copy the two times you got (congruent and incongruent) into appropriate columns.
- Don't forget to put your name there as well. (you can use a nickname, just don't leave the field empty)

`@hint`
- RTFM

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{r}

```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Thank you for particpating")
```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: f95c495704
xp: 50
```

- [Participate](http://www.onlinestrooptest.com/stroop_effect_test.php) in an online Stroop test.
- Open the [Google sheet](https://docs.google.com/spreadsheets/d/1IsyrFoKTCBpAOPnkdrwmHZnse6psvvZ1HCXyIaV50Jo/edit?usp=sharing) we will use to track you results.
- Copy the two times you got (congruent and incongruent) into appropriate columns.
- Don't forget to put your name there as well. (you can use a nickname, just don't leave the field empty)

Have you managed to do the test and write down the results?

`@hint`


`@possible_answers`
- [Yes]
- No

`@feedback`
- Great! Thank you!
- Don't sweat it. We probably got enough samples anyways.
